```python
# Initialize Otter
import otter
grader = otter.Notebook("week1.ipynb")
```

# HW1: Computational Notebooks and What is AI?
#### CS-31: Foundations for Artificial Intelligence
*Point breakdown: 40 autograded · 50 manually graded · 10 collaboration statement*

*You are allowed to work in small groups (2-3) for the lab. You must write anything that will be turned in -- all code and all written solutions -- on your own. You must list all people and resources that you interacted with in the Collaboration Statement below. See [course policies](https://www.cs.tufts.edu/comp/31/policy.html) for more information.*

--

In this assignment...

We are assuming you are comfortable with the following programming concepts:

- [Variables and elementary data types](https://nbviewer.org/github/jakevdp/WhirlwindTourOfPython/blob/master/03-Semantics-Variables.ipynb) (strings, integers, floats etc.)
- [Conditional expressions](https://nbviewer.org/github/jakevdp/WhirlwindTourOfPython/blob/master/07-Control-Flow-Statements.ipynb) (`if` and `else`)
- Logical Operators (AND and OR)
- [Loops](https://nbviewer.org/github/jakevdp/WhirlwindTourOfPython/blob/master/07-Control-Flow-Statements.ipynb) (`for` and `while`)
- [Basic collections](https://nbviewer.org/github/jakevdp/WhirlwindTourOfPython/blob/master/06-Built-in-Data-Structures.ipynb) (Lists and Dictionaries)
- Basic input/output (I/O)
- [Writing and calling functions](https://nbviewer.org/github/jakevdp/WhirlwindTourOfPython/blob/master/08-Defining-Functions.ipynb)
- Writing small (< 100 lines) programs that use the above concepts

We understand that you may not have programmed in a while, or that you've programmed before but not in Python. If you're new to Python, please take some time to read the linked parts of *A Whirlwind Tour of Python*. It will probably take time to remember some of the basics and get back in the groove, and that's ok! **However if after some time you find yourself severely struggling with the problems below, please see the instructor.**

### **Table of Contents**

0) Basic Notebook Navigation
1) Notebooks and Python
2) Python Warm Up Exercises
3) AI: The Big Picture
4) Intro Survey
5) Collaboration Statement

## **0. Basic Navigation and Autograder** 

Before we even start, doing these two things will help make navigating Jupyter Notebooks much easier and reduce clutter:

1) Near the top of this window, you'll see a button that says **Outline** (right next to **Jupyter Variables**). Click it.
    - You'll see a window pop up that shows a navigable outline for the notebook that you can keep handy.
2) Click on the text **0. Basic Navigation** right above this cell. You'll see a **⌄** icon right to the left. Click it, and it will collapse everything in this section. Right under it you'll see something like `1 cell hidden`. If you click it again you'll see a **>** symbol: click that again to expand the section.
    - Collapse all of the top level headers for sections 1, 2, and 3: and you can reveal them as you continue to keep the amount of text on the screen low.
    - You can do the same *within sections* (e.g. try collapsing and uncollapsing **0.1 Navigation Test** below) as you're working through them.


### **0.1 Navigation Test**

I'm hidden!

## **1. Notebooks and Python**

### **1.1 Cells**

Below is a *code cell*. These are where you do most of your computation.

If you click on the code cell, it will be highlighted on the left. You'll see a play button show up at the top left just outside of the cell, and two brackets (`[]`) show up below that. Sometimes you'll see this space to the left of the cell referred to as the *gutter*.

**Click the play icon to execute the cell**.

If you just opened up VSCode, you may see a window pop up near the top to select your Python environment. Select the CS31 environment. This is sometimes called the *kernel*.

When a cell is run, it automatically runs and then prints the last line in the cell. You'll also see a `[1]` show up in the bottom left. This is the *execution order* of the cell. You'll notice if you execute this cell multiple times in a row this will tick up.




```python
# Run me!

"Hello world!"
```

You can also run cells by highlighting it and pressing **Cmd/Ctrl+Enter** (my preferred method).

Running a cell evaluates all of the code inside of it.

```python
# Run me by pressing Cmd/Ctrl+Enter

9+10
```

If you want to run cells in a sequence, you can use **Shift+Enter** to run the cell and move focus to the next cell.

Try running the following cells in sequence using **Shift+Enter**.


```python
"Run me!"
```

```python
for i in range(1,10):
    print("Looping step", i)
```

```python
for name in ["James", "Harry", "Chris"]:
    print(f"Hello there {name}!")
```

### **1.2. Variables, Functions and Persistence**

As you may remember, variables allow us to store and manipulate data.

```python
my_number = 5

print(my_number)
```

Now, what's both cool and tricky about notebooks is that variables *persist* between cells. 

```python
your_number = my_number + 5

print(your_number)
print(my_number)
```

In the top bar, click on the icon that says `Jupyter Variables` (next to clear all outputs).

You should see a tab open up at the bottom that showcases all of the variables in memory.

What variables do you see? 

These are all the variables that are currently loaded into memory that you can reference.

### **1.3 Functions**

Just like in normal `.py` files, you can define functions and then use them throughout your notebook.

```python
def get_greater_than(input_list, number):
    new_list = []

    for item in input_list:
        if item > number:
            new_list.append(item)

    return new_list
```

```python
item = [1,10,2,9,3,8,4,7,5,6]

get_greater_than(item, 5)
```

#### **Exercise 1.1: Cells and Variables**

**It's time to do your first Exercise!**

Add a new cell which contains a function called `sum_up_numbers(start,end)` that takes in a `start` integer and an `end` integer and returns an integer that is the sum of all the integers from `start` to `end`, inclusive.

Then, below that cell, define the variable `sum_1_to_1000`. This variable should be an integer that holds the sum of all values from 1 to 1000, inclusive.

You can add a cell by pressing the `+ Code` button near the top. You can also select a cell and press `A` to add a cell above the currently selected cell, or `B` for below.

```python
def sum_up_numbers(start, end):
    ...
```

```python
...
```

```python
grader.check("q_sum_up_numbers")
```

### **1.4 Common Notebook Mistakes**

Having access to all your variables is nice, but it can lead to some interesting situations, such as this:

**Pause and Ponder**: What do you think will be the output if we run the following two cells in sequence?

```python
my_new_number = 10
your_new_number = my_new_number + 5


print(my_new_number)
print(your_new_number)
```

```python
my_new_number = 13
print(my_new_number)
print(your_new_number)
```

This also means that if you want to, it is allowable to run cells "out of order". 

Run the second cell and then the first cell.

```python
print(delete_me+1337)
```

```python
delete_me = 1337
```

**Common Jupyter Mistake #1**: You work on some analysis, delete some cells you were using for "scrap work", send the file over, and it doesn't work on your peer's machine because you forgot to delete some other cell.

**Before sharing (and submitting) a notebook, you should always `Restart`+`Run All` to make sure that it runs end-to-end with no issues!**

#### **Exercise 1.2: Delete Me!**

Delete the above two cells that contain `delete_me`, such that when you `Restart`+`Run All`, no errors occur. You can delete cells by clicking on them and clicking the trash icon in the top right of the cell, or highlighting a cell and pressing `D` twice.

**Common Notebook Mistake #2**: Changing Functions Without Rerunning

You change a function that you are using throughout your notebook, you run a cell that calls it, and it doesn't seem to have changed the result at all. 

**Always remember that when you change a function, you have to rerun it to confirm the change and overwrite the old version of the function that's in memory.** I've been using notebooks for probably around 8 years at this point and this is something I *still* forget to do sometimes.

As an example:

1) Execute the following two cells
2) Change filter_list to filter for numbers *less* than `number` **without** running the cell.
3) Run the second cell again - notice how the result does not change.
4) Run the `filter_list` cell again.
5) Then run the second cell again.

#### **Exercise 1.3: Fix `filter_list`.**

Once you've finished the walkthrough above, leave `filter_list` filtering for numbers **less than or equal to** `number` (the change from step 2) - that's the version that gets graded below.

```python
def filter_list(input_list, number):
    new_list = []

    for item in input_list:
        if item > number:
            new_list.append(item)

    return new_list
```

```python
list_to_filter = [1,2,3,4,5,6,7,8,9,10]

filter_list(list_to_filter, 4)
```

```python
grader.check("q_filter_list")
```

### **1.5 Markdown**

In addition to code cells, Jupyter also has *text* cells (which is what the instructions are!) These are useful for explaining in more detail your code, or putting figures or equations. These text cells are formatted using what's called [Markdown](https://www.markdownguide.org/basic-syntax/).

We write equations by using a formatting language called $\LaTeX$. It's what all of your professors use to write up your Math homework, and makes your documents look super professional. You write $\LaTeX$ by surrounding a math equation in dollar signs (\$).

For example, if I wanted to write $E=mc^2$, I would type `$E=mc^2$`. (Double click the cell to see this in action!)

#### **Exercise 1.4 (Written): Markdown Formatting.** **(5 points)**

**Welcome to your first written exercise!** Written questions will always be manually graded by the instructor or course staff. Written responses will be parsed out and graded as a separate assignment on Gradescope, so make sure you put your responses in the correct place. 

You're going to tell me about an equation! In the cell below, create a header with the name of your equation. Then, write the equation in $\LaTeX$. After that, provide two or three sentences on what the equation means and why it is important. In your explanation of the equations, some text should be **bolded** and some should be *italicized*.

You should reference [this page on how to write basic math expressions](https://www.overleaf.com/learn/latex/Mathematical_expressions). 

> **Note**: Sometimes I'll ask you to look things up online and even *\*gasp\** use AI Tools. I will preface this with a "**Look It Up**", which means overarching course rules of not searching for solutions or AI is temporarily suspended for that specified use case. Things I will ask you to look up are things that - in the context of this course - I don't consider core learning goals, but are good to have applied knowledge in.

> For example, it's not the goal of this course to make sure you know $\LaTeX$ commands by heart. This would be different if you were a graduate student or an industry researcher who is frequently writing reports and briefs that use mathematical notation, and thus it would be worth it to know it more deeply.

> **However, remember the overarching course rules. If I *don't* ask you to look something up, it is something that you should do on your own - as those are explicitly skills that I want you to build!**

**Look It Up**: Specialized math symbols can usually be found by searching online, and for this question you are also allowed to use AI Tools for translating screenshots of math to $\LaTeX$ code. Make sure it's correct!

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits. Your entire written answer must be contained in this cell._

<!-- END QUESTION -->

## **2. Python Warm Up Exercises**

If you need a refresher on basic Python syntax, or are coming from a different programming language, here are some helpful resources.

- [A Whirlwind Tour of Python](https://github.com/jakevdp/WhirlwindTourOfPython/) Chapters 1-9

### **Aside: Why Learn Programming in 2026?**


Some of you may have a nagging thought in the back of your head:
>*"Why even write code myself if I can ask a GenAI to generate it? I even know a bunch of professional programmers are using models to generate code, so why can't we?"*

Yes, you can make amazing things by programming with GenAI. I've done so myself! I'll even be frank, GenAI can do a lot of the *programming* questions (I've checked) quite well. 
 I could wax poetic all day about the intrinsic value of learning a skill, the joy of doing hard things, but then we would be here all day. Here's my pragmatic reason.

GenAI has been shown to create a *widening gap* in programming ability and software quality for those who use it.

For **novices that are still learning fundamental programming skills** who heavily used GenAI to solve problems, they often:
- Cannot explain the underlying logic or design choices of their program
- Have difficulty reasoning about what would happen if the underlying scenarios or context changed
- Can be misled by the AI into answering slightly different - but conceptually distinct - problems


In contrast, those with **strong foundational programming** skills are more often able to: 
- Discern when AI responses are actually useful and which responses are irrelevant
- Clearly communicate the logic of their program and how their design would change under different scenarios and context
- Know if the AI is making assumptions for them and is sending them down a wrong path

**You may notice that all of the underlying goals that the novices struggle with are goals of this course on understanding AI!**

Right now, it seems that the only proven method that has shown to increase the quality of people's skills in GenAI-assisted programming and *reasoning* about programs is **having good fundamentals in programming without GenAI!** Currently there's a lot of research and (heavy) debate on what exactly *strong foundational programming skills* are, and if we can teach high-quality GenAI-assisted development and computational reasoning without the traditional path of learning how to program. It might actually be the case that maybe in the future you can be a competent AI-powered programmer without learning the basics of programming first. That's something we probably won't know the answer to until 5-10 years from now. However, all the research suggests that people with strong fundamentals perform much better than those without, so instead of risking cheapening *your* long term effectiveness and knowledge for short term gains, we're gonna start from the basics.


Here's a [video explaining the above phenomenon](https://www.youtube.com/watch?v=HTUh0OO6Kmo), and the [research paper](https://dl.acm.org/doi/epdf/10.1145/3632620.3671116) that it was based on.



### **2.1 Reading and Debugging Python Code**



Knowing how to read and debug code is an important skill (*especially* with the rise of AI tools that generate large amounts of code), so we'll warm-up with Python with reading and debugging some existing code.

#### **Flower Gardens**

The function `biggest_bed` consumes a string (representing a garden) and produces a string representing the biggest bed of flowers in that string. You can see an example garden below. Notice how the garden is a single string with commas separating flowers of the same kind. Each group of flowers is a bed, and the goal of our function is to find the biggest (longest) bed. You can assume that all flowers in the bed are of the same type.



> \# Example Garden
>
>  "🌹🌹,🌷🌷,🌼🌼🌼,🌱🌱🌱🌱🌱,🌳,🌸,🌺🌺🌺🌺,🌻🌻🌻"


There are two extra difficulties to this function:

1. We need to ignore any beds that have seedlings ("🌱"), since those are not yet fully grown.
2.  We need to stop looking at any beds once we encounter a tree ("🌳"), since that marks the edge of our property.

In the example above, the answer would be "🌺🌺🌺🌺", were it not for these rules. Instead, the answer would become "🌼🌼🌼" since we ignore the row of seedlings and ignore everything past the tree.

The code below has a *broken* implementation of `biggest_bed`, which has decomposed the problem into various helper functions. There is one bug in each function. The bugs are all minor typos, and isolated to their own line. We might have changed a value, or deleted part of an expression, or moved variables around in a statement.

**We would recommend reading the `biggest_bed` function first to see the overall approach of the problem, before attempting to debug the individual subfunctions.** We also encourage you to add new cells as "scratch space" to view the output of the functions and create your own test examples.

#### **Some Recommendations and Tips Before Getting Started**
 - **DON'T JUMP RIGHT INTO WRITING/EDITING CODE**
 - Read all the code and subfunctions first before trying to debug individual functions
 - When making changes, write a little, then test a little. Don't try and do everything all at once.
 - When reading through code, write your own comments tracing your own understanding.

**Exercise 2.1: Debug `biggest_bed` and its subfunctions.**

```python
def take_until_trees(beds: list[str]) -> list[str]:
    taking = True
    result = []
    for bed in beds:
        if '🌳' in bed:
            taking = True
        elif taking:
            result.append(bed)
    return result
```

```python
grader.check("q_take_until_trees")
```

```python
def remove_seedlings(beds: list[str]) -> list[str]:
    result = []
    for bed in beds:
        if '🌱' in bed:
            result.append(bed)
    return result
```

```python
grader.check("q_remove_seedlings")
```

```python
def max_bed(beds: list[str]) -> str:
    most_popular = beds[0]
    for bed in beds:
        if len(bed) < len(most_popular):
            most_popular = bed
    return most_popular
```

```python
grader.check("q_max_bed")
```

```python
def biggest_bed(garden: str) -> str:
    if garden:
        return "no flowers"
    flower_beds = garden.split(',')
    flower_beds = take_until_trees(flower_beds)
    flower_beds = remove_seedlings(flower_beds)
    if not flower_beds:
        return 'no flowers'
    return max_bed(flower_beds)
```

```python
grader.check("q_biggest_bed")
```

### **2.2 Writing Code**

Now it's time to write some code on your own.

The local neighborhood is entering their gardens into a flower contest!

Let's say a _neighborhood_ is represented by a dictionary of strings. Each key is the owner of the garden, and the values are their gardens.

```python
# Example. Your functions should work on any dict in this format.
neighborhood = {
    "Raly": "🌹🌹,🌷🌷,🌼🌼🌼,🌳,🌸,🌺🌺🌺🌺,🌻🌻🌻",
    "Susan": "🌱🌱",
    "Crisp": "🌷🌷,🌻🌻🌻"
}
```

#### **Exercise 2.2: Neighborhood Winner**
Write a function called `neighborhood_winner` that takes in a neighborhood dictionary (as described above), and outputs the name of whoever has the biggest bed of flowers (no seedlings) within their property line. Assume no ties and that everyone has at least some flowers. 

In the given example neighborhood, `neighborhood_winner` would output `"Raly"`.

```python
def neighborhood_winner(neighborhood):
    ...
```

```python
grader.check("q_neighborhood_winner")
```

#### **Exercise 2.3: Flower Counter**

Write a function called `flower_counter(neighborhood, include_seeds)` which takes in a neighborhood dictionary, and outputs a dictionary where the keys are flowers, and the values are how many times that flower appears in the neighborhood. `include_seeds` is a boolean, that when set to True should include seeds in the final count - otherwise, the final count should ignore seeds. Make sure to follow property line rules!

In the given example neighborhood, `flower_counter(neighborhood, True)` would return a dictionary like:

`{'🌹': 2, '🌷': 4, '🌼': 3, '🌱': 2, '🌻': 3}`

Remember that order doesn't matter in dictionaries.


```python
def flower_counter(neighborhood, include_seeds):
    ...
```

```python
grader.check("q_flower_counter")
```

### **2.3 Plots**

For the last bit of programming, you should be somewhat familiar with how to make plots using `matplotlib`. We won't be doing too much with plots this week, but you'll see more types of plots in the future.

```python
import matplotlib.pyplot as plt
```

The most popular library for plotting in Python is named MatPlotLib. As a very large package, we only need to import a certain submodule within it. By convention and for convenience, we rename that module to be `plt`. You will see this line of code all the time at the start of programs in Python that make plots, so you should get used to typing it.

Here's how you make a simple line plot.

```python
data = [0, 5, 7, 35, 35, 51, 53, 57,
        58, 61, 64, 82, 84, 84, 88]

plt.plot(data)
plt.show()
```

Of course, a plot is no good without labels, so we can easily add those too.

```python
data = [0, 5, 7, 35, 35, 51, 53, 57,
        58, 61, 64, 82, 84, 84, 88]

plt.plot(data)

plt.xlabel("Time (Days)")
plt.ylabel("Money (Dollars)")
plt.title("Money over Time")

plt.show()
```

An example of a scatterplot.

```python
# Making a scatter plot

#        S1  S2  S3  S4  S5  S6  S7
exam1 = [40, 43, 55, 60, 52, 63, 74]
exam2 = [50, 55, 67, 72, 70, 71, 70]

plt.scatter(exam1, exam2)
plt.xlabel("Exam 1 Scores")
plt.ylabel("Exam 2 Scores")
plt.title("Exam Scores")
plt.show()
```

#### **Exercise 2.4: Flower Count Bar Plot** **(5 points)**

Make a bar chart that visualizes the results from a call to `flower_counter`, with the flower type on the x-axis and the count on the `y-axis`. 

`matplotlib` does not like having emojis as text, so we've provided you a dictionary that maps the emojis to their actual names. Make sure to give your plot descriptive labels and a title. 

- **Look It Up**: Look up how to make a bar plot in matplotlib! You'll notice that by default, the x-labels overlap each other and make the plot hard to read. Use documentation, a search engine, or an LLM to find ways to make the plot easier to read (e.g. rotating text labels, making text smaller, etc).

```python
flower_dict_names = {
    '🌹': "Rose", 
    '🌷': "Tulip", 
    '🌼': "Daisy", 
    '🌻': "Sunflower", 
    '🌸': "Cherry Blossom", 
    '🌱': "Seedling"
}
```

<!-- BEGIN QUESTION -->



```python
flower_dict = flower_counter(neighborhood, True)

...
```

<!-- END QUESTION -->

## **3. AI: The Big Picture**

#### **Exercise 3.1 (Written): Buzzword Bingo.** **(10 points)**

Find someone who is not studying AI or computing (e.g. a friend, a roommate, a family member) and ask them.

> *"How would you define artificial intelligence?"* 


If you can't find someone yourself, find someone's opinion (or "take") about AI online (e.g. a blog, a video, a social media post, etc).

In 3 to 5 sentences:
- Summarize their response
- Compare their definition to some of the definitions of AI that we talked about in class. Explicitly mention at least one of the definitions we talked about in class or present in the linked readings (such as Russell and Norvig's "thinking vs. acting" and "humanly vs. rationally"). Does their definition align with any of these, or is it something different? Did their definition surprise you in any way?

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

#### **Exercise 3.2 (Written): System Analysis.** **(10 points)**

Do some research and find a real-world tool, system, or object that you or other people consider (or once considered) "AI", "automatic", or "smart". Try and find something relevant to the field that you are currently studying. Unless your home field specializes in language and communication - try and find something *other* than a chatbot or Large Language Model!

Answer the following in 2 to 3 paragraphs.

- Give a short description of the system and what it is/was used for.
- What performance metric is the system designed to be evaluated on? If the system attempts to optimize multiple performance metrics at once, list them.
- Provide one possible *unintended consequence* of focusing solely on said performance measure(s). How could this be mitigated? (If it's a popular enough system, you may find real-life instances of an unintended consequence happening!)

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

#### **Exercise 3.3 (Written): "Why would it do that!?"** **(10 points)**

Listen to the podcast ["What Is AI, Anyway"](https://podcasts.apple.com/us/podcast/what-is-ai-anyway/id1749168252?i=1000658053334) featuring Peter Stone, a professor of CS at UT Austin.

Peter Stone says:

> "[AI] is a collection of technologies, it's not one thing... People tend to say 'Oh there's a computer that can beat the best chess player in the world, or the best go player in the world. That must mean that it's smarter than me. And if it's smarter than me, then it must be everything that I'm able to do.' "

Recall a time you were interacting with a "smart" or "intelligent" piece of technology (e.g. an app, a service that you use, etc) and it did something unexpected, unintended, or contrary to what you would expect an intelligent person would do. Summarize the encounter. Hypothesize one reason that might have caused a mismatch between your expectations and the technology's actual capability (e.g. marketing of the technology, poor instruction on how to use the software, confusing interface of the software, etc.).

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

#### **Exercise 3.4 (Written): What even is Intelligence?** **(10 points)**

Read [*Why aren't smart people happier?*](https://www.experimental-history.com/p/why-arent-smart-people-happier) by Adam Mastroianni and [*100,000,000 CROWPOWER and no horses on the moon*](https://taylor.town/crowpower) by Taylor Troesh. Both articles discuss the difficulties and history with defining and measuring intelligence. Write a short reflection (2-3 paragraphs) on the two essays. Be sure to answer the following questions.

- How does their discussion of comparing different tasks that require intelligence and how to measure them compare to the ways you've seen intelligence and capability measured in your own life? 
- Taylor ends with the claim that in the grand project of trying to understand intelligence that we should "measure anything". Do you agree with his conclusion? Are there some things we can't measure? Are there some things we *shouldn't* measure?
- How do these articles tie into your current understanding and perceptions of Artificial Intelligence?

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

## **4. Intro Survey**

Please remember to fill out the introductory survey, which is linked on the [schedule](https://www.cs.tufts.edu/cs/31/schedule.html) for Week 1. Your week 1 lab will not be graded until the survey is completed.

#### **Exercise 4.1: Confirm You Completed the Survey**

Change the following from `False` to `True` after you have completed the survey.


```python
completed_survey = False

```

```python
grader.check("q_completed_survey")
```

## **Collaboration Statement** **(10 points)**

You are required to complete the following collaboration statement.

<!-- BEGIN QUESTION -->

- If you completed the lab in a group - list all your group members here, including yourself. **If you worked individually, you must state *"I worked by myself."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- You must list the names of any others you got help from when completing this assignment (including TAs), and in what way you worked with them (discussed ideas, debugged math). **If you did not collaborate at all, you must state *"I did not collaborate with anyone else."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you found any resources online helpful, you must link them and describe how you found them helpful (or not helpful). **If you did not use any resources online, you must state *"I did not use any resources online."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you used any AI assistive technologies, you must describe how you used them and how you found them helpful (or not helpful). **If you did not use any AI assistive technologies, you must state *"I did not use any AI assistive technologies."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you have any questions or comments about the material, or there was topic or question you are confused on -- please make a note of it here! **If you do nothave any questions or comments, you must state *"I do not have any questions or comments about the assignment."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Run every cell from top to bottom and SAVE before exporting. Then submit
the generated .zip to corresponding assignment on Gradescope.


```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q_biggest_bed (points: None)
```
>>> assert biggest_bed('🌹🌹,🌷🌷,🌻🌻🌻,🌸') == '🌻🌻🌻'
```
```
>>> assert biggest_bed('🌹🌹,🌷🌷,🌼🌼🌼,🌸,🌺🌺🌺🌺,🌻🌻🌻') == '🌺🌺🌺🌺'
```
```
>>> assert biggest_bed('🌹🌹,🌷🌷,🌼🌼🌼,🌱🌱🌱🌱🌱,🌳,🌸,🌺🌺🌺🌺,🌻🌻🌻') == '🌼🌼🌼'
```
```
>>> assert biggest_bed('🌹🌹,🌷🌷,🌼🌼🌼,🌳,🌸,🌺🌺🌺🌺,🌻🌻🌻') == '🌼🌼🌼'
```
```
>>> assert biggest_bed('🌹🌹,🌷🌷,🌼🌼🌼,🌱🌱🌱🌱🌱,🌸,🌺🌺🌺🌺,🌻🌻🌻') == '🌺🌺🌺🌺'
```
```
>>> assert biggest_bed('🌱🌱,🌱🌱🌱,🌳,🌸,🌺🌺🌺🌺') == 'no flowers'
```
```
>>> assert biggest_bed('🌳,🌸,🌺🌺🌺🌺') == 'no flowers'
```
```
>>> assert biggest_bed('') == 'no flowers'
```
```
>>> assert biggest_bed('🌳🌳🌳') == 'no flowers'
```
```
>>> assert biggest_bed('🌱🌱') == 'no flowers'
```


Tests for q_completed_survey (points: None)
```
>>> assert completed_survey == True
```
failure_message: You must set completed_survey to True after completing the survey.


Tests for q_filter_list (points: None)
```
>>> assert filter_list(list_to_filter, 4) == [1, 2, 3, 4]
```
```
>>> assert filter_list([5, 10, 15], 10) == [5, 10]
```
```
>>> assert filter_list([3, 5, 5, 8, 1], 5) == [3, 5, 5, 1]
```


Tests for q_flower_counter (points: None)
```
>>> assert flower_counter(neighborhood, True) == {'🌹': 2, '🌷': 4, '🌼': 3, '🌱': 2, '🌻': 3}
```
```
>>> assert flower_counter(neighborhood, False) == {'🌹': 2, '🌷': 4, '🌼': 3, '🌻': 3}
```
```
>>> competition_neighborhood = {'Ming': '🌹🌹🌹,🌳,🌺🌺', 'Dana': '🌱🌱,🌻🌻🌻🌻', 'Ravi': '🌷🌷,🌼🌼🌼,🌸'}
>>> assert flower_counter(competition_neighborhood, True) == {'🌹': 3, '🌱': 2, '🌻': 4, '🌷': 2, '🌼': 3, '🌸': 1}
```
```
>>> competition_neighborhood = {'Ming': '🌹🌹🌹,🌳,🌺🌺', 'Dana': '🌱🌱,🌻🌻🌻🌻', 'Ravi': '🌷🌷,🌼🌼🌼,🌸'}
>>> assert flower_counter(competition_neighborhood, False) == {'🌹': 3, '🌻': 4, '🌷': 2, '🌼': 3, '🌸': 1}
```


Tests for q_max_bed (points: None)
```
>>> assert max_bed(['🌹🌹', '🌷🌷', '🌻🌻🌻', '🌸']) == '🌻🌻🌻'
```
```
>>> assert max_bed(['🌹🌹', '🌷🌷', '🌼🌼🌼', '🌸', '🌺🌺🌺🌺', '🌻🌻🌻']) == '🌺🌺🌺🌺'
```
```
>>> assert max_bed(['🌹🌹', '🌷🌷', '🌼🌼🌼', '🌱🌱🌱🌱🌱', '🌳', '🌸', '🌺🌺🌺🌺', '🌻🌻🌻']) == '🌱🌱🌱🌱🌱'
```
```
>>> assert max_bed(['🌺🌺🌺🌺', '🌱🌱', '🌱🌱🌱', '🌳', '🌸']) == '🌺🌺🌺🌺'
```
```
>>> assert max_bed(['🌹🌹🌹🌹', '🌼🌼', '🌸']) == '🌹🌹🌹🌹'
```


Tests for q_neighborhood_winner (points: None)
```
>>> assert neighborhood_winner(neighborhood) == 'Raly'
```
```
>>> competition_neighborhood = {'Ming': '🌹🌹🌹,🌳,🌺🌺', 'Dana': '🌱🌱,🌻🌻🌻🌻', 'Ravi': '🌷🌷,🌼🌼🌼,🌸'}
>>> assert neighborhood_winner(competition_neighborhood) == 'Dana'
```


Tests for q_remove_seedlings (points: None)
```
>>> assert remove_seedlings(['🌹', '🌱🌱', '🌷', '🌼🌼', '🌱🌱🌱', '🌳', '🌻🌻🌻']) == ['🌹', '🌷', '🌼🌼', '🌳', '🌻🌻🌻']
```
```
>>> assert remove_seedlings(['🌱🌱,🌱🌱🌱', '🌳', '🌸', '🌺🌺🌺🌺']) == ['🌳', '🌸', '🌺🌺🌺🌺']
```
```
>>> assert remove_seedlings(['🌷🌷', '🌱', '🌹', '🌼🌼']) == ['🌷🌷', '🌹', '🌼🌼']
```
```
>>> assert remove_seedlings(['🌱🌱']) == []
```


Tests for q_sum_up_numbers (points: None)
```
>>> assert sum_up_numbers(1, 10) == 55
```
```
>>> assert sum_1_to_1000 == 500500
```


Tests for q_take_until_trees (points: None)
```
>>> assert take_until_trees(['🌹🌹', '🌷🌷', '🌻🌻🌻']) == ['🌹🌹', '🌷🌷', '🌻🌻🌻']
```
```
>>> assert take_until_trees(['🌹🌹', '🌷🌷', '🌳', '🌻🌻🌻']) == ['🌹🌹', '🌷🌷']
```
```
>>> assert take_until_trees(['🌳', '🌸', '🌺🌺🌺🌺']) == []
```
```
>>> assert take_until_trees(['🌹', '🌳', '🌷🌷', '🌳', '🌸', '🌺🌺']) == ['🌹']
```
```
>>> assert take_until_trees(['🌳🌳🌳']) == []
```
